```python
# Initialize Otter
import otter
grader = otter.Notebook("randomStatisticsNumPy2.ipynb")
```

## Lecture Section

In this lecture, we will cover more of the `NumPy` library and the `statistics` module.
Specifically:
* NumPy Random
    * `np.random`
        * `.choice()`
        * `.randint()`
        * `.rand()`
        * `.shuffle()`
        * `.sample()`
    * Uniform Distribution
    * Normal Distribution
    * Exponential Distribution
* random module
    * When to use it...
* statistics module
    * When to use it...
    * `mean()`/`geometric_mean()`/`harmonic_mean()`
    * `median()`
    * `mode()`/`multimode()`
    * `variance()`/`stdev()`
    * `correlation()`/`covariance()`

### NumPy Random

Another powerful aspect of NumPy is its ability to create random numbers and distributions.

Run the lines below several times and observe what happens.

```python
import numpy as np
from numpy import random

random.randint(10)
```

```python
random.randint(10, size=(3, 2))
```

```python
random.rand()
```

```python
random.rand(3, 2)
```

We can randomly choose from a subset of items. Additionally, it has a parameter `p` where we can pass a list of probabilities for the choices.

```python
random.choice([1, 9, 3, 7], p = [0.90, 0.05, 0.02, 0.03])
```

We can use `random.shuffle()` to randomly chane the order of an array. This changes in-place, so it changes the original array. If we want to keep the original array the same, we can use `random.permutation()`

```python
arr = np.array([1, 2, 3, 4, 5, 6])
random.shuffle(arr)
arr
```

```python
arr = np.sort(arr)
arr_2 = random.permutation(arr)

print(arr)
print(arr_2)
```

Finally, NumPy's random lets you generate different distributions. I will show a few below, but there are many more.

```python
'''
Normal Distribution
loc = mean
scale = standard deviation
size = number of samples (or shape of returned array)
'''
random.normal(loc = 5, scale = 0.5, size = 12)
```

```python
'''
Uniform Distribution
low = lower bound
high = upper bound
size = number of samples (or shape of returned array)
'''
random.uniform(low = 0, high = 10, size = 12)
```

```python
'''
Exponential Distribution
scale = 1 / mean (inverse of rate)
size = number of samples (or shape of returned array)
'''
random.exponential(scale = 1, size = 12)
```

### `random` module

If you are familiar with Python, you may be wondering why we are using the `random` module from NumPy and not the base `random` module provided with Python.

In all honesty, they are not much different. The functions and abilities of both are similar, but the NumPy `random` nodule is better suited for statisticians, data scientists, and researchers because it is quicker, more efficient, and works with arrays.

The NumPy `random` module has a few extra distributions, and can sample items from an array much more efficiently than Python's `random`. We will use it over the base module, but the code block below shows you how to load in the base module for completion sake.

```python
import random
random.randint(5, 8)
```

### `statistics` module

In the next lesson, we will learn about a statistics package called `SciPy`. Before we get there, we will learn about the base `statistics` module. It has limited functionality—similar to a scientific calculator. However, it is still important. Unlike the two `random` modules introduced before, `SciPy` can not fully replace the `statistics` package.

Many of the functions mimic those in the NumPy package as well. When using arrays, we should use the NumPy package to use the same functions. Otherwise, the `statistics` package is valid.

One of the most important things to note about the `statistics` modules is that `NaN` or `Null`/`None` values cause undefined behavior. You should always strip them before using the `statistics` module.

```python
import statistics
data = [5, 6, 7, float('NaN')]
statistics.mean(data)
```

```python
from statistics import mean
data = [5, 6, 7]
mean(data)
```

We can also find the harmonic or geometric mean of a sequence.

```python
from statistics import harmonic_mean
harmonic_mean([1, 2, 3])

```

```python
from statistics import geometric_mean
geometric_mean([1, 2, 3])
```

We can grab the median of a sequence.

```python
from statistics import median
median([1, 2, 3])
```

We can also grab the median. `multimode()` wil get us the most common values of discrete data.

```python
from statistics import median
median([1, 2, 3, 3, 3, 4])
```

```python
from statistics import multimode
multimode(['blue', 'red', 'blue', 'green', 'red', 'red'])
```

We also have standard deviation and variance functions.

```python
from statistics import stdev
stdev([1, 3, 3, 9, 10])
```

```python
from statistics import variance
variance([1, 2, 3, 4, 5])
```

S.imilar to the `pandas` library, we also have a correlation and covariance function. As a quick reminder, correlation is a normalized covariance that ranges between -1 and 1.

```python
from statistics import correlation
x = [1, 2, 3, 4, 5, 6, 7, 8, 9]
y = [1, 2, 3, 1, 2, 3, 1, 2, 3]
correlation(x,y)
```

```python
from statistics import covariance
x = [1, 2, 3, 4, 5, 6, 7, 8, 9]
y = [1, 2, 3, 1, 2, 3, 1, 2, 3]
covariance(x, y)
```

Unfortunately, the statistics package won't work on anything other than lists. If you have a numpy dataframe and want to get the standard deviation of a column, converting with `tolist()` will NOT work. Instead, you need to use functions available to numpy, like `std()`.

## Assignment Section

**Question 1.**

We will borrow a dataset from Kaggle for this exercise. I have edited it for the purposes of this course.
https://www.kaggle.com/datasets/msjahid/colorado-motor-vehicle-sales-data

For this problem, you will use `pandas` to read the `"data/colorado_motor_vehicle_sales.csv"` file. Then, you will:
* Find the standard deviation of the `sales` column
* Find the mean of the `sales` column
* Create two new columns:
    * `sales_1std_down` which is the `sales` column minus one standard deviation
    * `sales_1std_above` which is the `sales` column plus one standard deviation

An example of the output is shown below. You can use the `head()` function to see the first few rows of the dataframe.

```
   year  quarter county        sales  sales_1sd_down  sales_1sd_above
0  2008        1  Adams  231609000.0    6.789792e+07     3.953201e+08
```

```python
import pandas as pd
from statistics import variance, stdev

df = ...
s_dev = ...
s_mean = ...
df['sales_1std_down'] = ...
df['sales_1std_above'] = ...


```

```python
grader.check("q1")
```

**Question 2.**

Using the dataframe from Question 1, find the most common county or counties in the `county` column. Save them in a list to `county_mode` - make sure it is alphabetical. Then. find how many unique counties there are and save the number to `unique_c`.

```python
from statistics import multimode
county_mode = ...
unique_c = ...
```

```python
grader.check("q2")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 6)
```
>>> assert np.isclose(df['sales_1std_down'].iloc[0], 68062704.94637245)
>>> assert np.isclose(df['sales_1std_above'].iloc[0], 395155295.05362755)
>>> assert np.isclose(df['sales_1std_down'].iloc[207], -134739295.05362755)
>>> assert np.isclose(df['sales_1std_above'].iloc[207], 192353295.05362755)
>>> assert np.isclose(s_dev, 163546295.05362755)
>>> assert np.isclose(s_mean, 171696919.51710263)
```


Tests for q2 (points: 4)
```
>>> assert unique_c == 18
>>> assert county_mode == ['Arapahoe', 'Denver', 'Douglas', 'El Paso', 'Fremont', 'Garfield', 'La Plata', 'Mesa']
>>> assert len(county_mode) == 8
```
